# LAB 2:  Campaign Financial Contributions [Total: 8 points]

## Instructions

The purpose of this assignment is for you to engage with a concrete data cleaning task. This will be accomplished through a coding assignment. You will carry out this task in the present notebook, and use the notebook to document the various steps of the exercise and to answer all questions.

In particular, we will be cleaning up some campaign donation data. The DC government makes data about campaign contribution (i.e. donations) available online to the general public on its website [opendata.dc.gov](https://opendata.dc.gov). The website does not provide very much documentation about the data.

We have created a smaller sample (`Campaign_Financial_Contributions_Sample.csv`) of this data for you to explore and clean. Follow the instructions in each question, and answer the accompanying questions. Note that some of the questions will require you to write a Python function to compute the correct answer.


## Important Warning
* Please ensure that you run the cell below before running any others. This will download all required files and packages, as well as installing the necessary packages to ensure the code runs successfully. If you restart the kernel or your runtime session (in Colab), be sure to rerun this cell before running any others.
* This assignment recommends using **Google Colab**. If you are using Anaconda Jupyter notebook/lab, please ensure that **this notebook is kept in a new folder**. This is because the following commands will delete all files with the extensions .csv and .py before downloading the required files.

```python
required_files = "https://github.com/mainuddin-rony/inst447-fall2024/raw/main/assignment/lab/lab2/required_files.zip"
! rm -rf tests
! rm -f required_files.zip *.csv *.py
! wget $required_files && unzip -j required_files.zip
! mkdir tests && mv *.py tests
! pip install otter-grader==5.5.0
```

Output:
```
zsh:1: no matches found: *.csv
--2024-10-04 10:26:33--  https://github.com/mainuddin-rony/inst447-fall2024/raw/main/assignment/lab/lab2/required_files.zip
Resolving github.com (github.com)... 140.82.114.3
Connecting to github.com (github.com)|140.82.114.3|:443... connected.
HTTP request sent, awaiting response... 302 Found
Location: https://raw.githubusercontent.com/mainuddin-rony/inst447-fall2024/main/assignment/lab/lab2/required_files.zip [following]
--2024-10-04 10:26:33--  https://raw.githubusercontent.com/mainuddin-rony/inst447-fall2024/main/assignment/lab/lab2/required_files.zip
Resolving raw.githubusercontent.com (raw.githubusercontent.com)... 185.199.109.133, 185.199.110.133, 185.199.108.133, ...
Connecting to raw.githubusercontent.com (raw.githubusercontent.com)|185.199.109.133|:443... connected.
HTTP request sent, awaiting response... 200 OK
Length: 80959 (79K) [application/zip]
Saving to: ‘required_files.zip’

required_files.zip  100%[===================>]  79.06K  --.-KB/s 
```

```python
# Initialize Otter
import otter
grader = otter.Notebook()
```

```python
import pandas as pd
```

# Part A: Inspecting the data frame

The first set of questions will ask you to inspect the data. Feel free to add new code cells and make sure to comment your code with markdown cells to describe your process of reaching the answer.

## Q1 (Manual Grading)

**Points:** 1

**Question:** Read the sample campaign contribution data into a pandas data frame. Inspect the columns and cells of the data set. What is most likely the column name of the variable which has the method of payment in which the donation/contribution was made?

1. `GIS_LAST_MOD_DTTM`

2. `DATEOFRECEIPT`

3. `CONTRIBUTIONTYPE`

Write down the correct answer in the following cell.

**`CONTRIBUTIONTYPE`**

```python
df = pd.read_csv("Campaign_Financial_Contributions_Sample.csv")
df.head()
```

Output:
```
Unnamed: 0  OBJECTID                              COMMITTEENAME  \
0      205514    205514                                 Fenty 2006   
1      169351    168351                    Muriel Bowser for Mayor   
2      129869    129869                    Muriel Bowser for Mayor   
3      112855    112855  Washington DC Association of Realtors PAC   
4       88950     88950  Washington DC Association of Realtors PAC   

   CANDIDATENAME  ELECTIONYEAR       CONTRIBUTORNAME  \
0   Adrian Fenty          2006    Stefon  Brathwaite   
1  Muriel Bowser          2014        Denise Rollins   
2  Muriel Bowser          2014  Antoinette Barksdale   
3            NaN          2002        Alethea Jordan   
4            NaN          2002        Matthew Rogers   

                                       ADDRESS CONTRIBUTORTYPE  \
0           2141 I St NW, Washington, DC 20037      Individual   
1    3649 11th Street NW, Washington, DC 20010      Individual   
2        3626 Park Pl NW, Washington, DC 20010 
```

```python
df.columns
```

Output:
```
Index(['Unnamed: 0', 'OBJECTID', 'COMMITTEENAME', 'CANDIDATENAME',
       'ELECTIONYEAR', 'CONTRIBUTORNAME', 'ADDRESS', 'CONTRIBUTORTYPE',
       'CONTRIBUTIONTYPE', 'EMPLOYER', 'EMPLOYERADDRESS', 'AMOUNT',
       'DATEOFRECEIPT', 'ADDRESS_ID', 'XCOORD', 'YCOORD', 'LATITUDE',
       'LONGITUDE', 'FULLADDRESS', 'GIS_LAST_MOD_DTTM', 'WARD', 'Unnamed: 21',
       'Unnamed: 22', 'Unnamed: 23', 'Unnamed: 24', 'Unnamed: 25',
       'Unnamed: 26', 'Unnamed: 27', 'Unnamed: 28', 'Unnamed: 29',
       'Unnamed: 30', 'Unnamed: 31', 'Unnamed: 32', 'Unnamed: 33',
       'Unnamed: 34', 'Unnamed: 35', 'Unnamed: 36', 'Unnamed: 37',
       'Unnamed: 38', 'Unnamed: 39', 'Unnamed: 40'],
      dtype='object')
```

## Q2 (Auto Grading)

**Points:** 1

**Question:** How many donations (rows) in this dataset are from corporations?

Write a function called `numdoncorp()` the function should take as parameter the name of the file with data and it should return the answer to the question as an int.

```python
import pandas as pd

def numdoncorp(path):
  df = pd.read_csv(path)
  return df[df['CONTRIBUTORTYPE'] == 'Corporation'].shape[0]
```

```python
numdoncorp("Campaign_Financial_Contributions_Sample.csv")
```

Output:
```
137
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Q3 (Manual Grading)

**Points:** 1

**Question:** True or False: The campaign contribution data set is structured data.

1. True
2. False

Write down the answer in the following cell.

**`TRUE`**

# Part B: cleaning inconsistent records

The next set of questions will ask you to clean the dataframe before computing the desired value.

## Q4 (Auto Grading)

**Points:** 1

**Question:**  What is the minimum donation amount (do not round)? Note that in this an the following questions, any donation amounts of 0 are likely mistakes and are NOT acceptable answers because donors cannot give a $0 donation. Since there are some donations with 0 amount, filter these rows.

Write a function called `findLeastAmount()`. the function should take as parameter the name of the file with data and it should return the exact value and not round off value as the answer.

```python
import pandas as pd

def findLeastAmount(path):
  df = pd.read_csv(path)
  return df[df['AMOUNT'] > 0]['AMOUNT'].min()
```

```python
findLeastAmount("Campaign_Financial_Contributions_Sample.csv")
```

Output:
```
0.09
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## Q5 (Auto Grading)

<div class="alert alert-warning">Make sure to apply the same cleaning step from Q4</div>

**Points:** 1

**Question:** What is the maximum donation amount?

Write a function called `findMaxAmount()`. the function should take as parameter the name of the file with data and it should return the exact value and not round off value as the answer.

```python
import pandas as pd

def findMaxAmount(path):
  df = pd.read_csv(path)
  return df[df['AMOUNT'] > 0]['AMOUNT'].max()
```

```python
findMaxAmount("Campaign_Financial_Contributions_Sample.csv")
```

Output:
```
35000.0
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

## Q6 (Auto Grading)

<div class="alert alert-warning">Make sure to apply the same cleaning step from Q4</div>

**Points:** 1

**Question:** What is the most recent date listed in the column for DATEOFRECEIPT?

Write a function called `findRecentDate()`. the function should take as parameter the name of the file with data and it should return a Python string in the form `YYYY-MM-DD`.

```python
import pandas as pd

def findRecentDate(path):
  df = pd.read_csv(path)
  df['DATEOFRECEIPT'] = pd.to_datetime(df['DATEOFRECEIPT'])
  return df['DATEOFRECEIPT'].max().strftime('%Y-%m-%d')
```

```python
findRecentDate("Campaign_Financial_Contributions_Sample.csv")
```

Output:
```
'2019-03-13'
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

## Q7 (Auto Grading)

<div class="alert alert-warning">Make sure to apply the same cleaning step from the Q4</div>

**Points: 1**

**Question:** How many missing values are there for contributor type (CONTRIBUTORTYPE)?

```python
import pandas as pd

def findMissingValues(path):
  df = pd.read_csv(path)
  return df[df['AMOUNT'] > 0]['CONTRIBUTORTYPE'].isnull().sum()
```

```python
findMissingValues("Campaign_Financial_Contributions_Sample.csv")
```

Output:
```
1
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

## Submission

Don't forget to run all cells in your notebook and then save it. To save, click on `File`, then select `Save/Save Notebook`. After that, download the notebook by going to `File`, then `Download` (for Anaconda notebook), and choosing `Download .ipynb` (for Colab). Finally, submit the notebook on Gradescope using the link found on ELMS.

Tests for q2 (points: None)
```
>>> assert numdoncorp('Campaign_Financial_Contributions_Sample.csv') == 137, "The number of rows doesn't match with expected result"
```


Tests for q4 (points: None)
```
>>> assert findLeastAmount('Campaign_Financial_Contributions_Sample.csv') == 0.09, "The return value doesn't match with the expected value"
```


Tests for q5 (points: None)
```
>>> assert findMaxAmount('Campaign_Financial_Contributions_Sample.csv') == 35000.0, "The return value doesn't match with the expected value"
```


Tests for q6 (points: None)
```
>>> assert findRecentDate('Campaign_Financial_Contributions_Sample.csv') == '2019-03-13', "The return date doesn't match with the expected date"
```


Tests for q7 (points: None)
```
>>> assert findMissingValues('Campaign_Financial_Contributions_Sample.csv') == 1, "The return value doesn't match with the expected value"
```
